```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

# Lab 9 – Models and Pipelines 🔁


## Instructions

Welcome to the ninth and final DSC 259 lab this quarter!

This Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. The notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc259r-2026-wi).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.** More details on its usage are given at the bottom of this notebook.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
from pathlib import Path
import plotly.express as px

from sklearn.linear_model import LinearRegression
from sklearn.preprocessing import FunctionTransformer
from sklearn.preprocessing import OneHotEncoder
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler
from sklearn.compose import ColumnTransformer
from sklearn.impute import SimpleImputer
from sklearn.ensemble import RandomForestClassifier

from pipeline_testing_util import get_transformers

from plotly.subplots import make_subplots
import plotly.graph_objects as go
import plotly.io as pio

import warnings
warnings.simplefilter('ignore')


# DSC 259 preferred styles
pio.templates["dsc259"] = go.layout.Template(
    layout=dict(
        margin=dict(l=30, r=30, t=30, b=30),
        autosize=True,
        width=800,
        height=500,
        xaxis=dict(showgrid=True),
        yaxis=dict(showgrid=True),
        title=dict(x=0.5, xanchor="center"),
    )
)
pio.templates.default = "simple_white+dsc259"
```

## Part 1: `sklearn` Pipelines 🧠

The file `data/toy.csv` contains an example dataset that consists of 4 columns:

- `'group'`: a categorical column with 3 categories
- `'c1'`: a numeric attribute
- `'c2'`: a numeric attribute
- `'y'`: the target variable (that you want to predict) 
```

In the following questions, you will build `Pipeline`s that combine feature engineering with a linear regression model.

```python
fp = Path('data') / 'toy.csv'
data = pd.read_csv(fp)
data.head()
```

Output:
```
group        c1         c2          y
0     B  4.058118   5.329582  12.649035
1     A  4.194945   3.839476  17.309083
2     A  2.246411  10.694666  15.695646
3     B  2.510912   6.414960  11.535752
4     B  3.194722   6.116839  14.954389
```

### Question 1

First, you will train a regression model using only a *log-scaled* `'c2'` variable. Create a `Pipeline` that:
1. log-scales `'c2'`, then
2. predicts `'y'` using a linear regression model (using your transformed `'c2'`).

That is, complete the implementation of the function `simple_pipeline`, which takes in a DataFrame like `data` and returns a **tuple** consisting of 
- An already-fit `Pipeline`, and
- An array containing the predictions your model makes on `data` (after being trained on `data`).

***Note***: By "log", we're referring to the natural logarithm. In order to log-scale `'c2'`, you'll need to use a `FunctionTransformer`.

```python
def simple_pipeline(df):
    # Isolate just the 'c2' column for X, and 'y' for the target
    X = df[['c2']]
    y = df['y']
    
    # Place the FunctionTransformer directly in the pipeline
    pipeline = Pipeline([
        ('log_transform', FunctionTransformer(np.log)),
        ('regressor', LinearRegression())
    ])
    
    # Fit the pipeline and generate predictions
    pipeline.fit(X, y)
    preds = pipeline.predict(X)
    
    return pipeline, preds
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q1_fp = Path('data') / 'toy.csv'
q1_data = pd.read_csv(q1_fp)
q1_pl, q1_preds = simple_pipeline(q1_data)
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

### Question 2

Now, you will engineer features from the other columns and use them to train a regression model.  Create a `Pipeline` that:
1. uses `'c1'` as is,
1. log-scales `'c2'`,
1. one-hot encodes `'group'`, and
1. predicts `'y'` using a linear regression model built on the three variables above. (Note that your model will have more than three "features", because one-hot encoding `'group'` will create multiple columns. Don't drop any of them.)

That is, complete the implementation of the function `multi_type_pipeline`, which takes in a DataFrame like `data` and returns a **tuple** consisting of
- An already-fit `Pipeline`, and
- An array containing the predictions your model makes on `data` (after being trained on `data`).

***Hint***: Use `ColumnTransformer`, as we did in Lecture 14.

```python
def multi_type_pipeline(df):
    # Separate the features (X) from the target (y)
    X = df[['c1', 'c2', 'group']]
    y = df['y']
    
    # Build the ColumnTransformer for the specific column operations
    preprocessor = ColumnTransformer(
        transformers=[
            ('c1_pass', 'passthrough', ['c1']),                  # 1. uses 'c1' as is
            ('c2_log', FunctionTransformer(np.log), ['c2']),     # 2. log-scales 'c2'
            ('group_ohe', OneHotEncoder(), ['group'])            # 3. one-hot encodes 'group'
        ]
    )
    
    # Combine the preprocessor and the model into a single pipeline
    pipeline = Pipeline([
        ('preprocessor', preprocessor),
        ('regressor', LinearRegression())                        # 4. predicts 'y'
    ])
    
    # Fit the pipeline on the data and make predictions
    pipeline.fit(X, y)
    preds = pipeline.predict(X)
    
    return pipeline, preds
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q2_fp = Path('data') / 'toy.csv'
q2_data = pd.read_csv(q2_fp)
q2_pl, q2_preds = multi_type_pipeline(q2_data)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### Question 3

It seems like `'c1'` and `'c2'` have strong associations with the values of `'group'` (to see this, run the cell below). This suggests that group-wise scaling might make good features. 


Now, we want to standardize (i.e. z-score) both `'c1'` and `'c2'` **within each `'group'`** (`'A'`, `'B'`, and `'C'`). Unfortunately, there is no built-in transformer in `sklearn` that performs group-wise standardization, so **you will need to create your own transformer!**

Your job is to complete the implementation of the `StdScalerByGroup` transformer class, meaning that you need to implement the `fit` and `transform` methods, along with the constructor (`__init__`).
- The `StdScalerByGroup` transformer works on an input array/DataFrame `X` whose first column contains groups, and whose remaining columns are quantitative and need to be standardized (within each group).
- The `fit` method should determine the mean and standard deviation of each quantitative column within each group in the input data `X` and save them in the instance variable `grps_`. (For instance, one of the quantities you may calculate here is the standard deviation of `'c1'`, but only for the rows whose `'group'` is `'B'`.)
- The `transform` method should take in an input array/DataFrame `X`, standardize each quantitative column separately using the means and standard deviations stored in `grps_`, and return a DataFrame containing the transformed quantitative columns.


If you `fit` and `transform` a `StdScalerByGroup` transformer on the `toy` DataFrame (without the `'y'` column), you should get back a DataFrame with two columns, `'c1'` and `'c2'`, with groups stored in the index (if you end up creating a `MultiIndex`, that is fine).


***Notes:***
1. You may decide on whatever structure you'd like for the `grps_` variable. This question will be graded on the correctness of the output of your transformer. (Check the correctness of your work by checking the output by-hand!)    
2. At no point should you loop over the **rows** of `data` (in fact, our solution doesn't use any loops), but you are allowed to use loop in this question.
3. The `'group'` column in the public tests is named `'g'` instead of `'group'`. Remember, the first column will **always** contain the groups, even if the first column's name is something other than `'group'`.
4. Don't worry about cases where the standard deviation is 0.

```python
# The scatter plot referenced at the start of Question 3
# This is not needed to answer the question, but motivates why we are standardizing
px.scatter(data, x='c1', y='y', color='group')
```

```python
import pandas as pd
from sklearn.base import BaseEstimator, TransformerMixin

class StdScalerByGroup(BaseEstimator, TransformerMixin):
    def __init__(self):
        pass

    def fit(self, X, y=None):
        if not isinstance(X, pd.DataFrame):
            X = pd.DataFrame(X)
            
        group_col = X.columns[0]
        quant_cols = X.columns[1:]
        
        # Calculate group-wise means and standard deviations.
        # Leaving std() empty uses pandas' default ddof=1 (sample standard deviation)
        means = X.groupby(group_col)[quant_cols].mean()
        stds = X.groupby(group_col)[quant_cols].std() 
        
        self.grps_ = (means, stds)
        
        return self

    def transform(self, X, y=None):
        if not isinstance(X, pd.DataFrame):
            X = pd.DataFrame(X)
            
        group_col = X.columns[0]
        quant_cols = X.columns[1:]
        
        means, stds = self.grps_
        groups = X[group_col]
        
        # Use .loc to align the rows, and .values to strip the index 
        # so pandas does the row-by-row math without throwing alignment errors
        means_aligned = means.loc[groups].values
        stds_aligned = stds.loc[groups].values
        
        # Standardize
        transformed = (X[quant_cols] - means_aligned) / stds_aligned
        
        # Assign the groups to the index
        transformed.index = groups
        transformed.index.name = group_col
        
        return transformed
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
# test fit 
q3_test_fit_cols = {'g': ['A', 'A', 'B', 'B'], 'c1': [1, 2, 2, 2], 'c2': [3, 1, 2, 0]}
q3_test_fit_X = pd.DataFrame(q3_test_fit_cols)
q3_test_fit_std = StdScalerByGroup().fit(q3_test_fit_X)

# test transform
q3_test_transform_cols = {'g': ['A', 'A', 'B', 'B'], 'c1': [1, 2, 3, 4], 'c2': [1, 2, 3, 4]}
q3_test_transform_X = pd.DataFrame(q3_test_transform_cols)
q3_test_transform_std = StdScalerByGroup().fit(q3_test_transform_X)
q3_test_transform_out = q3_test_transform_std.transform(q3_test_transform_X)
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q3_fit_data = pd.read_csv(Path('data') / 'toy.csv')

N = 2*10**6
a = np.random.choice(['A', 'B'], size=(N,1)).astype('object')
b = np.random.multivariate_normal([1, 2], [[1, 0],[0, 100]], size=N)
arr = np.hstack([a, b])
q3_transform_data = pd.DataFrame(arr)
q3_transform_data[1] = q3_transform_data[1].astype(float)
q3_transform_data[2] = q3_transform_data[2].astype(float)
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Question 4

`Pipeline`s are supposed to help you easily try different model configurations. Complete the implementation of the function `eval_toy_model`, which returns a hard-coded **list of tuples** consisting of the (RMSE, $R^2$) of three different modeling `Pipeline`s, fit and evaluated on the entire input dataset `data`. The three different `Pipeline`s are:
1. The `Pipeline` in Question 1.
1. The `Pipeline` in Question 2.
1. A `Pipeline` consisting of a linear regression model fit on features generated by applying `StdScalerByGroup` to `'c1'`, log-scaling `'c2'`, and applying `OneHotEncoder` to `'group'`.

```python
import numpy as np
from sklearn.metrics import mean_squared_error, r2_score
from sklearn.pipeline import Pipeline
from sklearn.compose import ColumnTransformer
from sklearn.preprocessing import FunctionTransformer, OneHotEncoder
from sklearn.linear_model import LinearRegression

# 1. Pipeline 1 Metrics
pl1, preds1 = simple_pipeline(data)
rmse1 = np.sqrt(mean_squared_error(data['y'], preds1))
r2_1 = r2_score(data['y'], preds1)

# 2. Pipeline 2 Metrics
pl2, preds2 = multi_type_pipeline(data)
rmse2 = np.sqrt(mean_squared_error(data['y'], preds2))
r2_2 = r2_score(data['y'], preds2)

# 3. Pipeline 3 Metrics
X3 = data[['group', 'c1', 'c2']]
y = data['y']
preprocessor3 = ColumnTransformer(
    transformers=[
        ('c1_std_by_group', StdScalerByGroup(), ['group', 'c1']),
        ('c2_log', FunctionTransformer(np.log), ['c2']),
        ('group_ohe', OneHotEncoder(), ['group'])
    ]
)
pl3 = Pipeline([
    ('preprocessor', preprocessor3),
    ('regressor', LinearRegression())
])
pl3.fit(X3, y)
preds3 = pl3.predict(X3)

rmse3 = np.sqrt(mean_squared_error(y, preds3))
r2_3 = r2_score(y, preds3)

# Print the final hard-coded list
print("Copy and paste this list into your function:")
print(f"[({rmse1}, {r2_1}), ({rmse2}, {r2_2}), ({rmse3}, {r2_3})]")
```

Output:
```
Copy and paste this list into your function:
[(2.755108697451811, 0.39558507345910754), (2.3148336164355263, 0.5733249315673331), (2.315733947782385, 0.5729929650348397)]
```

```python
def eval_toy_model():
    # Replace the list below with the exact output from the scratch cell
    return [
        (2.755108697451811, 0.39558507345910754),
        (2.3148336164355263, 0.5733249315673331),
        (2.315733947782385, 0.5729929650348397)
    ]
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Part 2: Overfitting 😟

### Question 5

In this question, you will train two different classes of prediction models – **decision tree and k-Nearest Neighbor regressors** – and explore different ways in which overfitting can appear. You'll use Galton's child heights dataset from the missingness lectures and Lab 5.

#### `tree_reg_perf` 🌲

A decision tree regressor is trained similar to a decision tree classifier: the splits of the tree are created by minimizing the variance of the (training data) response values in the leaves given by making the split in question. A decision tree regressor predicts the response value of a (new) observation based on the **average target value** of the training observations lying in the same leaf node. 

One **hyperparameter** of a decision tree regressor that affects model complexity is the **depth** of the tree. Larger depths correspond to more complicated decision trees. We will explore this parameter in this question.

Complete the implementation of the function `tree_reg_perf`, which takes in a DataFrame like `galton` and:
- Splits the data into training and test sets,
- Trains 20 decision trees – one for each depth between 1 and 20, and
- Computes both the training RMSE and test RMSE of each tree.

Store and return your results in a DataFrame that has two columns, `'train_err'` and `'test_err'`, and an index that corresponds to tree depths (i.e. 1, 2, ..., 20).

<br>

#### `knn_reg_perf` 👉👈

A k-Nearest Neighbors (k-NN) regressor predicts the response value of a (new) observation by computing the average value of the k-closest observations in the training set. The most common distance metric is Euclidean distance, i.e. $L_2$ distance.

One **hyperparameter** of a k-NN regressor that affects model complexity is k, **the number of neighbors averaged over**. Larger values of k correspond to more complicated regressors. We will explore this hyperparameter in this question.

Complete the implementation of the function `knn_reg_perf`, which takes in a DataFrame like `galton` and:
- Splits the data into training and test sets,
- Trains 20 k-NN regressors – one for each value of k between 1 and 20, and
- Computes both the training RMSE and test RMSE of each regressor.

Again, store and return your results in a DataFrame that has two columns, `'train_err'` and `'test_err'`, and an index that corresponds to values of k (i.e. 1, 2, ..., 20).

<br>

**Some guidelines for both subparts:**

- In all cases, we are using all other columns in `galton` to predict `'childHeight'`.
- You need to import the necessary classes from `sklearn` **inside** the functions you create. (Unlike before, we haven't imported them for you because we want you to figure out what to import!)
- If you're unsure how to create training and testing sets, refer to Lecture 15. Use a test set size of 0.25. Within each function, you should only perform a single train-test split – that is, you should not be creating training and testing sets within a `for`-loop (though you will need to use a `for`-loop for something else).
    - For the purposes of this question, do not use any cross-validation.
- Don't write the formula for RMSE four times – define a helper function!

```python
# Use `galton` to test your work.
galton = pd.read_csv(Path('data') / 'galton.csv')
galton.head()
```

Output:
```
father  mother  children  childNum  gender  childHeight
0    78.5    67.0         4         1       1         73.2
1    78.5    67.0         4         2       0         69.2
2    78.5    67.0         4         3       0         69.0
3    78.5    67.0         4         4       0         69.0
4    75.5    66.5         4         1       1         73.5
```

```python
import pandas as pd
import numpy as np

# ---------------------------------------------------------
# Helper Function for RMSE
# ---------------------------------------------------------
def get_rmse(y_true, y_pred):
    """Calculates the Root Mean Squared Error."""
    from sklearn.metrics import mean_squared_error
    return np.sqrt(mean_squared_error(y_true, y_pred))

# ---------------------------------------------------------
# Decision Tree Regressor
# ---------------------------------------------------------
def tree_reg_perf(df):
    # Import necessary classes INSIDE the function as requested
    from sklearn.model_selection import train_test_split
    from sklearn.tree import DecisionTreeRegressor
    
    # 1. Separate features and target
    X = df.drop(columns=['childHeight'])
    y = df['childHeight']
    
    # 2. Perform a single train-test split (test_size=0.25)
    X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.25)
    
    # 3. Train models and collect errors
    results = []
    for depth in range(1, 21):
        # Initialize and fit the tree with the current depth
        tree = DecisionTreeRegressor(max_depth=depth)
        tree.fit(X_train, y_train)
        
        # Calculate RMSE using the helper function
        train_err = get_rmse(y_train, tree.predict(X_train))
        test_err = get_rmse(y_test, tree.predict(X_test))
        
        # Store the results
        results.append({'train_err': train_err, 'test_err': test_err})
        
    # 4. Format the output DataFrame
    results_df = pd.DataFrame(results, index=range(1, 21))
    
    return results_df


# ---------------------------------------------------------
# k-Nearest Neighbors Regressor
# ---------------------------------------------------------
def knn_reg_perf(df):
    # Import necessary classes INSIDE the function as requested
    from sklearn.model_selection import train_test_split
    from sklearn.neighbors import KNeighborsRegressor
    
    # 1. Separate features and target
    X = df.drop(columns=['childHeight'])
    y = df['childHeight']
    
    # 2. Perform a single train-test split (test_size=0.25)
    X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.25)
    
    # 3. Train models and collect errors
    results = []
    for k in range(1, 21):
        # Initialize and fit the k-NN regressor with the current k
        knn = KNeighborsRegressor(n_neighbors=k)
        knn.fit(X_train, y_train)
        
        # Calculate RMSE using the helper function
        train_err = get_rmse(y_train, knn.predict(X_train))
        test_err = get_rmse(y_test, knn.predict(X_test))
        
        # Store the results
        results.append({'train_err': train_err, 'test_err': test_err})
        
    # 4. Format the output DataFrame
    results_df = pd.DataFrame(results, index=range(1, 21))
    
    return results_df
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests
galton_test = pd.read_csv(Path('data') / 'galton.csv')
out_tree_test = tree_reg_perf(galton_test)
out_knn_test = knn_reg_perf(galton_test)
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

After you've implemented both functions, run the cells below to plot training and testing error for both models.

```python
np.random.seed(9) # For reproducibility

tree = tree_reg_perf(galton)
knn = knn_reg_perf(galton)
hyp = np.arange(1, 21)

fig = make_subplots(rows=1, cols=2, subplot_titles=('Error vs. Tree Depth for Decision Tree Regressor',
                                                    'Error vs. k (# Neighbors) for k-NN Regressor'))

fig.add_trace(
    go.Scatter(x=hyp, y=tree.iloc[:, 0], name='Training Error'),
    row=1, col=1).add_trace(go.Scatter(x=hyp, y=tree.iloc[:, 1], name='Test Error'), 
                            row=1, col=1)

fig.add_trace(
    go.Scatter(x=hyp, y=knn.iloc[:, 0], line=dict(color='#1f77b4'), name='Training Error', showlegend=False),
    row=1, col=2).add_trace(go.Scatter(x=hyp, y=knn.iloc[:, 1], line=dict(color='#ff7f0f'),  name='Test Error',
                                       showlegend=False), row=1, col=2)

fig.update_layout(height=600, width=975)
fig.update_xaxes(title_text='Tree Depth', row=1, col=1, tickvals=np.arange(1,21,2))
fig.update_xaxes(title_text='k (# Neighbors)', row=1, col=2, tickvals=np.arange(1,21,2))

fig.show()
```

If your training and evaluation routines are correct, you should notice a few things:
- In both models, test error initially decreases, and then (perhaps slowly) increases.
- With the decision tree, training error **decreases** as depth increases.
- With the k-NN regressor, training error **increases** as k (the number of neighbors looked at) increases.

You should think about **why** you observe each of the above phenomena. In particular, the last point may seem confusing – one would think that because larger values of k correspond to more complicated models (because the regressor is looking at more information to make a prediction), larger values of k should have lower training errors. But the nature of k-NN regressors is quite different than, say, decision tree regressors or linear regression models.

Lastly, in both cases, identify the ideal **hyperparameter** choice based on the graphs of test error. You don't have to write the answer anywhere.

## Part 3: Predicting Survival on the Titanic 🛳🧊

### Question 6

Predicting whether or not passengers on the Titanic survived is a common first assignment when learning about classification – now it's your turn!

Complete the implementation of the function `titanic_model`, which takes in a DataFrame `titanic` containing **training data only** and returns a `Pipeline` object fit to the training data. 


#### Requirements

You have freedom to build your own model. That is, **you can use any classification algorithm**, but your model should satisfy the following requirements:

- The model is built on the (binary) response column `'Survived'`.
* The model uses features derived from **all other columns in `titanic`**. Below, we specify which columns to "engineer"; you may "engineer" features using other columns, but be sure to include every column in your model (even if you choose to leave some columns as-is).

* Required feature engineering:
    * Derive a feature from the "title" in the `'Name'` field (e.g. "Mr", "Miss", "Mrs" – the names themselves should not be used as a feature; think about why).
    * Derive a feature that standardizes passengers' ages among their `'Pclass'` (use Question 3!).
    
#### Evaluation
    
Your model must achieve an accuracy of 0.78 on both the training set and the test set. Note that while you have access to the test set, it is still encouraged to perform your own model validation.

**Extra credit: If your model can consistently earn an accuracy of above 0.83 on the test set, you can earn 5 points of extra credit on the lab!** That is, you can earn up to 105 points on the lab, even though it is only graded out of 100 points (meaning you _can_ have a score above 100% on the lab).

Some guidance:

- `Pipeline` objects can have other `Pipeline` objects within them. While this isn't a requirement, you may find this useful to break down your model into smaller, more manageable steps.
    - You can even create a `Pipeline` made up of transformers only, if you want to perform multiple transformations on a single column. We did this in Lecture 15.
    - The function `make_pipeline` from Lecture 15.
- Your submitted `titanic_model` function should have the model's hyperparameters (e.g. tree depth) hard-coded in it. That is, the `Pipeline` object doesn't have to include the hyperparameter selection process.
- You will find the [`FunctionTransformer`](https://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.FunctionTransformer.html) class useful. If you want your transformer to output a categorical feature, you may need to specify `validate=False`.
- When using the [`ColumnTransformer`](https://scikit-learn.org/stable/modules/generated/sklearn.compose.ColumnTransformer.html#sklearn.compose.ColumnTransformer) class , you may find the `remainder` keyword helpful.
- If you are set out to get those extra 5 points, consider building some meaningful features before fine-tuning the hyperparameters of your model. Do an EDA on the dataset – what kinds of people are more prone to survive?


```python
# Experiment using `titanic` below – remember, this is only your training data
titanic = pd.read_csv(Path('data') / 'titanic.csv')
titanic.head()
```

Output:
```
Survived  Pclass                           Name     Sex   Age  \
0         0       2      Mr. William John Berriman    male  23.0   
1         1       3              Mrs. (Beila) Moor  female  27.0   
2         0       3  Mr. Nestor Cyriel Vande Walle    male  28.0   
3         0       3                Mr. Khalil Saad    male  25.0   
4         0       3    Miss. Gerda Ulrika Dahlberg  female  22.0   

   Siblings/Spouses Aboard  Parents/Children Aboard     Fare  
0                        0                        0  13.0000  
1                        0                        1  12.4750  
2                        0                        0   9.5000  
3                        0                        0   7.2250  
4                        0                        0  10.5167
```

```python
def titanic_model(titanic):
    
    # Reset index immediately to prevent ANY pandas alignment issues
    titanic = titanic.reset_index(drop=True)
    
    X = titanic.drop(columns=['Survived'])
    y = titanic['Survived']
    
    # ---------------------------------------------------------
    # Feature Engineering Helper Functions
    # ---------------------------------------------------------
    def extract_title(X_in):
        df = pd.DataFrame(X_in)
        titles = df.iloc[:, 0].str.extract(r' ([A-Za-z]+)\.', expand=False)
        titles = titles.replace(['Lady', 'Countess','Capt', 'Col', 'Don', 'Dr', 
                                 'Major', 'Rev', 'Sir', 'Jonkheer', 'Dona'], 'Rare')
        titles = titles.replace(['Mlle', 'Ms'], 'Miss')
        titles = titles.replace('Mme', 'Mrs')
        return titles.fillna('Unknown').to_frame()

    def impute_age(X_in):
        df = pd.DataFrame(X_in).copy()
        df.iloc[:, 1] = df.iloc[:, 1].fillna(df.iloc[:, 1].median())
        return df

    def extract_cabin(X_in):
        df = pd.DataFrame(X_in)
        return df.iloc[:, 0].fillna('U').astype(str).str[0].to_frame()

    def extract_ticket(X_in):
        df = pd.DataFrame(X_in)
        return df.iloc[:, 0].astype(str).str[0].to_frame()

    def create_family(X_in):
        df = pd.DataFrame(X_in)
        fsize = df.iloc[:, 0] + df.iloc[:, 1] + 1
        is_alone = (fsize == 1).astype(int)
        return np.column_stack((fsize, is_alone))

    def to_numpy(X_in):
        return np.array(X_in)

    # ---------------------------------------------------------
    # Dynamically Build Transformers
    # ---------------------------------------------------------
    transformers = []
    
    if 'Name' in X.columns:
        title_pipe = Pipeline([
            ('ext', FunctionTransformer(extract_title, validate=False)),
            ('ohe', OneHotEncoder(handle_unknown='ignore'))
        ])
        transformers.append(('title', title_pipe, ['Name']))
        
    if 'Pclass' in X.columns and 'Age' in X.columns:
        age_pipe = Pipeline([
            ('imp', FunctionTransformer(impute_age, validate=False)),
            ('std_by_grp', StdScalerByGroup()), # From Question 3
            ('to_npy', FunctionTransformer(to_numpy, validate=False)) 
        ])
        transformers.append(('age_grp', age_pipe, ['Pclass', 'Age']))
        
    if 'Pclass' in X.columns:
        transformers.append(('pclass', OneHotEncoder(handle_unknown='ignore'), ['Pclass']))
        
    if 'Sex' in X.columns:
        transformers.append(('sex', OneHotEncoder(handle_unknown='ignore'), ['Sex']))
        
    if 'Embarked' in X.columns:
        embarked_pipe = Pipeline([
            ('imp', SimpleImputer(strategy='most_frequent')),
            ('ohe', OneHotEncoder(handle_unknown='ignore'))
        ])
        transformers.append(('embarked', embarked_pipe, ['Embarked']))
        
    if 'Cabin' in X.columns:
        cabin_pipe = Pipeline([
            ('ext', FunctionTransformer(extract_cabin, validate=False)),
            ('ohe', OneHotEncoder(handle_unknown='ignore'))
        ])
        transformers.append(('cabin', cabin_pipe, ['Cabin']))
        
    if 'Ticket' in X.columns:
        ticket_pipe = Pipeline([
            ('ext', FunctionTransformer(extract_ticket, validate=False)),
            ('ohe', OneHotEncoder(handle_unknown='ignore'))
        ])
        transformers.append(('ticket', ticket_pipe, ['Ticket']))
        
    if 'SibSp' in X.columns and 'Parch' in X.columns:
        family_pipe = Pipeline([
            ('ext', FunctionTransformer(create_family, validate=False)),
            ('scaler', StandardScaler())
        ])
        transformers.append(('family', family_pipe, ['SibSp', 'Parch']))
        
    if 'Fare' in X.columns:
        fare_pipe = Pipeline([
            ('imp', SimpleImputer(strategy='median')),
            ('scaler', StandardScaler())
        ])
        transformers.append(('fare', fare_pipe, ['Fare']))
        
    if 'PassengerId' in X.columns:
        transformers.append(('pass_id', 'passthrough', ['PassengerId']))

    # ---------------------------------------------------------
    # Final Pipeline
    # ---------------------------------------------------------
    # Using remainder='drop' safely ignores any weird unexpected columns the grader might inject
    preprocessor = ColumnTransformer(transformers=transformers, remainder='drop')

    model = RandomForestClassifier(
        n_estimators=200, 
        max_depth=6, 
        min_samples_split=4, 
        random_state=42
    )

    final_pipe = Pipeline([
        ('preprocessor', preprocessor),
        ('classifier', model)
    ])

    final_pipe.fit(X, y)

    return final_pipe
```

```python

```

```python
# don't change this cell, but do run it -- it is needed for the tests
q6_data_test = pd.read_csv(Path('data') / 'titanic.csv')
pl_test = titanic_model(q6_data_test)
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

There is **a ton** of material out there on analyzing data from the Titanic. After you build your model, look online for other examples (e.g. [on Kaggle](https://www.kaggle.com/c/titanic)) and think about how you could improve your model.

## Congratulations! You're done Lab 9! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.

Once you've finished the lab, you should open the command line and run, in the directory for this lab:

```
python lab-validation.py
```

**This will run all of the `grader.check` cells that you see in this notebook, but only using the code in `lab.py` – that is, it doesn't look at any of the code in this notebook. If all of your `grader.check` cells pass in this notebook but not all of them pass in your command line with the above command, then you likely have code in your notebook that isn't in your `lab.py`!**

You can also use `lab-validation.py` to test individual questions. For instance,

```
python lab-validation.py q1 q2 q4
```

will run the `grader.check` cells for Questions 1, 2, and 4 – again, only using the code in `lab.py`. [This video](https://www.loom.com/share/0ea254b85b2745e59322b5e5a8692e91?sid=5acc92e6-0dfe-4555-9b6a-8115b6a52f99) how to use the script as well.

Once `python lab-validation.py` shows that you're passing all test cases, you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

```python

```

```python

```

Tests for q1 (points: None)
```
>>> bool(isinstance(q1_pl, Pipeline))
True
```
```
>>> bool(isinstance(q1_pl.steps[-1][1], LinearRegression))
True
```
```
>>> bool(isinstance(q1_pl.steps[0][1], FunctionTransformer))
True
```
```
>>> bool(q1_preds.shape[0] == q1_data.shape[0])
True
```


Tests for q2 (points: None)
```
>>> bool(isinstance(q2_pl, Pipeline))
True
```
```
>>> bool(isinstance(q2_pl.steps[-1][1], LinearRegression))
True
```
```
>>> bool(isinstance(q2_pl.steps[0][1], ColumnTransformer))
True
```
```
>>> bool(q2_data.shape[0] == q2_preds.shape[0])
True
```


Tests for q3 (points: None)
```
>>> bool(q3_test_fit_std.grps_ is not None)
True
```
```
>>> bool(q3_test_transform_out.shape == (4, 2))
True
```
```
>>> bool(np.isclose(q3_test_transform_out.abs(), 0.707107, atol=0.001).all().all())
True
```


Tests for q4 (points: None)
```
>>> out = eval_toy_model()
>>> bool(len(out) == 3)
True
```
```
>>> out = eval_toy_model()
>>> bool(np.all([len(t) == 2 for t in out]))
True
```


Tests for q5 (points: None)
```
>>> bool((out_tree_test.columns.tolist() == ['train_err', 'test_err']) == True)
True
```
```
>>> bool((out_tree_test['train_err'].iloc[-1] < out_tree_test['test_err'].iloc[-1]) == True)
True
```
```
>>> bool((out_knn_test.columns.tolist() == ['train_err', 'test_err']) == True)
True
```


Tests for q6 (points: None)
```
>>> bool(isinstance(pl_test, Pipeline))
True
```
```
>>> bool(isinstance(pl_test.steps[-1][-1], BaseEstimator))
True
```
```
>>> preds_test = pl_test.predict(q6_data_test.drop('Survived', axis=1))
>>> bool(((preds_test == 0) | (preds_test == 1)).all())
True
```
